<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 84424  -->
  <question type="coderunner">
    <name>
      <text>PROTOTYPE_python3_parsons_problem</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<h2>The prototype for a Parson's Problem in Python</h2>
<p>This prototype defines, via the CodeRunner <code>__twigprefix__</code> mechanism a single Twig macro <code>global_extra(jumbled_code)</code> which takes as a parameter a multiline string of left-justified Python3 program lines.</p><p>Questions of this type must set their global extra field to the single line</p>
  <pre>  {% _self.global_extra(jumbled_code) %}  </pre>
<p>where <code>jumbed_code</code> is a Twig variable. It will normally be defined via a template parameter to be a multiline string obtained by taking the Python program (the goal program that the student is expected to achieve), stripping all indentation and shuffling the order of the lines.
</p><p>For example, if a Python3 template pre-processor is used, the template parameters for a question of this type might be of the form:</p>
<pre>import json, random
code = '''def print_squares(nums):
    """Print the squares of numbers in the list nums"""
    i = 0
    while i &lt; len(nums):
       print(nums[i] ** 2)
       i += 1
'''
code_lines = [line.strip() for line in code.splitlines()]
random.shuffle(code_lines)
answer = {"CR-parsons-code": [code]}
print(json.dumps({"jumbled_code": '\n'.join(code_lines),
                  "answer": json.dumps(answer)}))
</pre>
<p>which also defines a template parameter for the answer, so that the question author's answer field can simple be written as <code>{{answer}}</code>.</p>
<p>The question text and all the test data are just as normal for testing the code defined as the <code>code</code> variable within the template parameter.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3_parsons_problem</coderunnertype>
    <prototypetype>2</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>5</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace>1</useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import json
SEPARATOR = "#<ab@17943918#@>#"
student_code = json.loads("""{{ STUDENT_ANSWER | e('py') }}""")["CR-parsons-code"][0]
exec(student_code)

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins>0</allowmultiplestdins>
    <answer></answer>
    <validateonsave>0</validateonsave>
    <testsplitterre><![CDATA[|#<ab@17943918#@>#\n|ms]]></testsplitterre>
    <language>python3</language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>NearEqualityGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams><![CDATA[import json
__twigprefix__ = r"""{% macro global_extra(jumbled_code) %}
<!--
Code for use within the CodeRunner HTML UI to implement a Parsons
Problem question type.

Much of this code is derived from the jsparson project by Ihantola
and Karavirta, licensed as below.

-- Richard Lobb, August 2021.

jsparsons licence:
=================
The MIT License

Copyright (c) 2010- Petri Ihantola and Ville Karavirta

Permission is hereby granted, free of charge, to any person obtaining a copy
of this software and associated documentation files (the "Software"), to deal
in the Software without restriction, including without limitation the rights
to use, copy, modify, merge, publish, distribute, sublicense, and/or sell
copies of the Software, and to permit persons to whom the Software is
furnished to do so, subject to the following conditions:

The above copyright notice and this permission notice shall be included in
all copies or substantial portions of the Software.

THE SOFTWARE IS PROVIDED "AS IS", WITHOUT WARRANTY OF ANY KIND, EXPRESS OR
IMPLIED, INCLUDING BUT NOT LIMITED TO THE WARRANTIES OF MERCHANTABILITY,
FITNESS FOR A PARTICULAR PURPOSE AND NONINFRINGEMENT. IN NO EVENT SHALL THE
AUTHORS OR COPYRIGHT HOLDERS BE LIABLE FOR ANY CLAIM, DAMAGES OR OTHER
LIABILITY, WHETHER IN AN ACTION OF CONTRACT, TORT OR OTHERWISE, ARISING FROM,
OUT OF OR IN CONNECTION WITH THE SOFTWARE OR THE USE OR OTHER DEALINGS IN
THE SOFTWARE.
-->
<style>
/** Stylesheet for the puzzles */

.sortable-code {
    position: static;
    padding-left: 0px;
    margin-left: 2%;
    float: left;
    width: 94%;
    max-width: 700px;
}

.sortable-code ul {
    font-size: 90%;
    font-family: monospace;
    list-style: none;
    background-color: #efefff;
    padding-bottom: 10px;
    padding-left: 0;
    margin-left: 0;
    border: 1px solid #efefff;;
}
.sortable-code ul:empty {
  padding-bottom: 30px;
}
.sortable-code li, .sortable-code li:before, .sortable-code li:after {
  box-sizing: content-box;
}
ul.output {
    background-color: #FFA;
}
.sortable-code li {
    -moz-border-radius:10px;
    -webkit-border-radius:10px;
    border-radius: 10px;
    background-color: white;
    border:1px solid lightgray;
    padding:4px;
    padding-left:8px;
    margin-top: 4px;
    white-space: nowrap;
    overflow: hidden;
    cursor: move;
}
.sortable-code li:hover {
    overflow: visible;
}


</style>

<script>
require(['jquery', 'jqueryui'], function($, jqui) { 

    function turn_on_touch() {
        // Turn on touch functionality
        // Detect touch support

        /*!
        * jQuery UI Touch Punch 0.2.3
        *
        * Copyright 2011–2014, Dave Furfero
        * Dual licensed under the MIT or GPL Version 2 licenses.
        *
        * Depends:
        *  jquery.ui.widget.js
        *  jquery.ui.mouse.js
        */
        $.support.touch = 'ontouchend' in document;

        // Ignore browsers without touch support
        if (!$.support.touch) {
            return;
        }

      var mouseProto = $.ui.mouse.prototype,
          _mouseInit = mouseProto._mouseInit,
          _mouseDestroy = mouseProto._mouseDestroy,
          touchHandled;

        /**
        * Simulate a mouse event based on a corresponding touch event
        * @param {Object} event A touch event
        * @param {String} simulatedType The corresponding mouse event
        */
        function simulateMouseEvent (event, simulatedType) {

        // Ignore multi-touch events
        if (event.originalEvent.touches.length > 1) {
          return;
        }

        event.preventDefault();

        var touch = event.originalEvent.changedTouches[0],
            simulatedEvent = document.createEvent('MouseEvents');

        // Initialize the simulated mouse event using the touch event's coordinates
        simulatedEvent.initMouseEvent(
          simulatedType,    // type
          true,             // bubbles                    
          true,             // cancelable                 
          window,           // view                       
          1,                // detail                     
          touch.screenX,    // screenX                    
          touch.screenY,    // screenY                    
          touch.clientX,    // clientX                    
          touch.clientY,    // clientY                    
          false,            // ctrlKey                    
          false,            // altKey                     
          false,            // shiftKey                   
          false,            // metaKey                    
          0,                // button                     
          null              // relatedTarget              
        );

        // Dispatch the simulated event to the target element
        event.target.dispatchEvent(simulatedEvent);
        }

        /**
        * Handle the jQuery UI widget's touchstart events
        * @param {Object} event The widget element's touchstart event
        */
        mouseProto._touchStart = function (event) {

        var self = this;

        // Ignore the event if another widget is already being handled
        if (touchHandled || !self._mouseCapture(event.originalEvent.changedTouches[0])) {
          return;
        }

        // Set the flag to prevent other widgets from inheriting the touch event
        touchHandled = true;

        // Track movement to determine if interaction was a click
        self._touchMoved = false;

        // Simulate the mouseover event
        simulateMouseEvent(event, 'mouseover');

        // Simulate the mousemove event
        simulateMouseEvent(event, 'mousemove');

        // Simulate the mousedown event
        simulateMouseEvent(event, 'mousedown');
        };

        /**
        * Handle the jQuery UI widget's touchmove events
        * @param {Object} event The document's touchmove event
        */
        mouseProto._touchMove = function (event) {

        // Ignore event if not handled
        if (!touchHandled) {
          return;
        }

        // Interaction was not a click
        this._touchMoved = true;

        // Simulate the mousemove event
        simulateMouseEvent(event, 'mousemove');
        };

        /**
        * Handle the jQuery UI widget's touchend events
        * @param {Object} event The document's touchend event
        */
        mouseProto._touchEnd = function (event) {

        // Ignore event if not handled
        if (!touchHandled) {
          return;
        }

        // Simulate the mouseup event
        simulateMouseEvent(event, 'mouseup');

        // Simulate the mouseout event
        simulateMouseEvent(event, 'mouseout');

        // If the touch interaction did not move, it should trigger a click
        if (!this._touchMoved) {

          // Simulate the click event
          simulateMouseEvent(event, 'click');
        }

        // Unset the flag to allow other widgets to inherit the touch event
        touchHandled = false;
      };

      /**
       * A duck punch of the $.ui.mouse _mouseInit method to support touch events.
       * This method extends the widget with bound touch event handlers that
       * translate touch events to mouse events and pass them to the widget's
       * original mouse event handling methods.
       */
      mouseProto._mouseInit = function () {
        
        var self = this;

        // Delegate the touch handlers to the widget's element
        self.element.bind({
          touchstart: $.proxy(self, '_touchStart'),
          touchmove: $.proxy(self, '_touchMove'),
          touchend: $.proxy(self, '_touchEnd')
        });

        // Call the original $.ui.mouse init method
        _mouseInit.call(self);
      };

      /**
       * Remove the touch event handlers
       */
      mouseProto._mouseDestroy = function () {
        
        var self = this;

        // Delegate the touch handlers to the widget's element
        self.element.unbind({
          touchstart: $.proxy(self, '_touchStart'),
          touchmove: $.proxy(self, '_touchMove'),
          touchend: $.proxy(self, '_touchEnd')
        });

        // Call the original $.ui.mouse destroy method
        _mouseDestroy.call(self);
      };
    }; 
    
    turn_on_touch(); // Run the jquery ui touch code.
    
    // =====================================================
    //
    // Now the code based on jsparsons (see licence at top)
    // ----------------------------------------------------
    // =====================================================

    var MAX_INDENT = 7;
    
    // regexp used for trimming
    var trimRegexp = /^\s*(.*?)\s*$/;
    var translations = {
        en: {
            code_panel_label: 'Use drag and drop to reorder and indent the following code lines.',
        }
    };

    // =================================================================
    //
    // ParsonsCodeLine object definition
    // ---------------------------------
    //
    // Defines a line object skeleton with only code and indentation from
    // a code string of the problem definition string (see parseCode)
    // =================================================================
    var ParsonsCodeLine = function(codestring, widget, allow_indent) {
        this.widget = widget;
        this.code = "";
        this.indent = 0;
        if (codestring) {
            // Consecutive lines to be dragged as a single block of code have strings "\\n" to
            // represent newlines => replace them with actual new line characters "\n"
            this.code = codestring.replace(/#distractor\s*$/, "").replace(trimRegexp, "$1").replace(/\\n/g, "\n");
            if (allow_indent) {
                this.indent = (codestring.length - codestring.replace(/^\s+/, "").length) / 4;
                this.indent = Math.min(this.indent, MAX_INDENT);
            }
            this.id = widget.id_prefix + ParsonsCodeLine.next_id++;
        }
    };

    ParsonsCodeLine.next_id = 0; // Effectively a static variable of the 'class'.


    // =================================================================
    //
    // ParsonsWidget object definition
    // -------------------------------
    // Creates a parsons widget.
    // Init must be called after creating an object.
    //
    // ==================================================================
    var ParsonsWidget = function(options) {
        var defaults = {
            'incorrectSound': false,
            'x_indent': 35,
            'feedback_cb': false,
            'lang': 'en'
        };

        this.options = jQuery.extend({}, defaults, options);
        this.id_prefix = options.sortableId + 'codeline';
        if (translations.hasOwnProperty(this.options.lang)) {
            this.translations = translations[this.options.lang];
        } else {
            this.translations = translations['en'];
        }

        // translate code_panel_label
        if (!this.options.hasOwnProperty("code_panel_label")) {
            this.options.code_panel_label = this.translations.code_panel_label;
        }
    };

    // Parse a problem definition given as a string and returns 
    // a list of ParsonCodeLines.
    ParsonsWidget.prototype.parseCode = function(lines, allow_indent) {
        var lineList = [],
            lineObject,
            that = this;
        // Create line objects out of each codeline.
        // Fields in line objects:
        //   code: a string of the code, may include newline characters and 
        //     thus in fact represents a block of consecutive lines
        //   indent: indentation level.
        $.each(lines, function(index, item) {
            lineObject = new ParsonsCodeLine(item, that, allow_indent);
            // Initialize line object with code and indentation properties
            if (lineObject.code.length > 0) {
                // The line is non-empty, not just whitespace
                lineList.push(lineObject);
            }
        });

        return lineList;
    };

    ParsonsWidget.prototype.init = function(code_string) {
        // Set up the widget, given the code to be displayed as a
        // single string. 
        // Can't use jQuery to get element directly as it can't parse
        // the element ID properly due to ___textareaId___ macro.
        var lines = this.parseCode(code_string.split("\n"), true);
        var that = this;
        var sortableElement = document.getElementById(this.options.sortableId);
        var sortable = $(sortableElement);
        html = '<p>' + this.options.code_panel_label + '</p>' +
            this.codeLinesToHTML(lines, this.options.sortableId);
        
        sortable.html(html);

        this.manageDrags();
    };


    ParsonsWidget.prototype.calcIndent = function(leftDiff) {
        // Return the indent based on the pixel difference given.
        return Math.max(0, Math.round(leftDiff / this.options.x_indent));
    };

    ParsonsWidget.prototype.getCodeString = function(element_id) {
        // Extract and return the code from the DOM as a single string.
        var parent = document.getElementById(element_id);
        var codeLines = parent.childNodes;
        var code = "";
        codeLines.forEach(item => {
            var line = item.innerText.trim();
            var indent = parseInt(item.style.marginLeft, 10);
            indent = indent / this.options.x_indent;
            var newLine = " ".repeat(indent * 4) + line + '\n';
            code += newLine;
        });
        return code;
    };

    ParsonsWidget.prototype.getStudentCodeString = function() {
        // Return the lines of code as a single string.
        return this.getCodeString('ul-' + this.options.sortableId);
    };

    ParsonsWidget.prototype.displayError = function(message) {
        if (this.options.incorrectSound && $.sound) {
            $.sound.play(this.options.incorrectSound);
        }
        alert(message);
    };

    ParsonsWidget.prototype.setHTMLIndent = function(codelineID, indent) {
        var codeLine = $(document.getElementById(codelineID));
        codeLine.css("margin-left", this.cssIndent(indent));
    };

    ParsonsWidget.prototype.cssIndent = function(indent) {
        return this.options.x_indent * Math.min(MAX_INDENT, indent) + "px";
    };

    ParsonsWidget.prototype.codeLineToHTML = function(codeline) {
        var style = codeline.indent ? ' style="margin-left:' + this.cssIndent(codeline.indent) + '" ' : '';
        return '<li id="' + codeline.id + '" class="ace-highlight-code lang-py"' + style + '>' + codeline.code + '<\/li>';
    };


    ParsonsWidget.prototype.codeLinesToHTML = function(codeLines, destinationID) {
        // Construct and return a single <ul> element with the HTML
        // for the given list of lines as embedded LI elements. The UL
        // element's ID is destinationID.
        var lineHTML = [];
        codeLines.forEach(line => {
            lineHTML.push(this.codeLineToHTML(line));
        });
        return '<ul id="ul-' + destinationID + '">' + lineHTML.join('') + '</ul>';
    };

    ParsonsWidget.prototype.manageDrags = function() {
        // Set up the code pane as a sortable UI elements with drag
        // management for indentation. An annoying issue here is that
        // the ___textareaId___ macro produces IDs that jQuery cannot
        // handle, so we have to use getElementById everywhere.
        var that = this;
        var sortableElement = document.getElementById('ul-' + this.options.sortableId);
        var sortable = $(sortableElement).sortable({
            start: function(event, ui) {
             },
            stop: function(event, ui) {
                if ($(event.target)[0] != ui.item.parent()[0]) {
                    return;
                }
                // The helper isn't where it seems to be because of the margin!
                var currentMargin = parseInt(ui.item.css('margin-left'));
                var realPosition = currentMargin + ui.position.left - ui.item.parent().position().left;
                var indent = that.calcIndent(realPosition);
                that.setHTMLIndent(ui.item[0].id, indent);
            },
            grid: [this.options.x_indent, 1]
        });
        sortable.addClass("output");
    };
    
   window.ParsonsWidget = ParsonsWidget;
   
    function displayErrors(fb) {
        if(fb.errors.length > 0) {
            alert(fb.errors[0]);
        }
    } 

    $(document).ready(function() {
        var sortableId = 'sortable-___textareaId___';
        var crTextareaId = 'CR-parsons-code-___textareaId___';
        var parson = new ParsonsWidget({
            'sortableId': sortableId
        });

        var codeTextArea = document.getElementById(crTextareaId);

        var code = codeTextArea.value; // Extract the code as a single multiline string.
        parson.init(code);
        
        function sync_to_coderunner(mutations) {
            codeTextArea.value = parson.getStudentCodeString();
        };
        dropArea=document.getElementById('ul-' + sortableId);
        let observer=new MutationObserver(sync_to_coderunner);
        let observerOptions = {
            childList: true,
            attributes: true,
            characterData: false,
            subtree: false,
            attributeFilter: ['style'],
            attributeOldValue: false,
            characterDataOldValue: false
        };
        observer.observe(dropArea, observerOptions);
        if (window.ace && window.applyAceHighlighting) {
            window.applyAceHighlighting(window.ace, dropArea);
        }
    });
});
</script>

<html>
    <head>
        <title>Simple js-parsons example</title>
    </head>
    <body>
        <div id="sortable-___textareaId___" class="sortable-code">
        </div>
        <div style="clear:both;"></div>
        <div id="CRactual"  style="display: none;">
        <h2>Code Here:</h2>
<textarea id="CR-parsons-code-___textareaId___" name='CR-parsons-code' class="coderunner-ui-element" rows="10" cols="40">
{{ jumbled_code }}
</textarea>
        </div>
    </body>
</html>

{% endmacro %}
"""
print(json.dumps({'__twigprefix__': __twigprefix__}))
]]></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>python3</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald><![CDATA[{"__twigprefix__": "{% macro global_extra(jumbled_code) %}\n<!--\nCode for use within the CodeRunner HTML UI to implement a Parsons\nProblem question type.\n\nMuch of this code is derived from the jsparson project by Ihantola\nand Karavirta, licensed as below.\n\n-- Richard Lobb, August 2021.\n\njsparsons licence:\n=================\nThe MIT License\n\nCopyright (c) 2010- Petri Ihantola and Ville Karavirta\n\nPermission is hereby granted, free of charge, to any person obtaining a copy\nof this software and associated documentation files (the \"Software\"), to deal\nin the Software without restriction, including without limitation the rights\nto use, copy, modify, merge, publish, distribute, sublicense, and/or sell\ncopies of the Software, and to permit persons to whom the Software is\nfurnished to do so, subject to the following conditions:\n\nThe above copyright notice and this permission notice shall be included in\nall copies or substantial portions of the Software.\n\nTHE SOFTWARE IS PROVIDED \"AS IS\", WITHOUT WARRANTY OF ANY KIND, EXPRESS OR\nIMPLIED, INCLUDING BUT NOT LIMITED TO THE WARRANTIES OF MERCHANTABILITY,\nFITNESS FOR A PARTICULAR PURPOSE AND NONINFRINGEMENT. IN NO EVENT SHALL THE\nAUTHORS OR COPYRIGHT HOLDERS BE LIABLE FOR ANY CLAIM, DAMAGES OR OTHER\nLIABILITY, WHETHER IN AN ACTION OF CONTRACT, TORT OR OTHERWISE, ARISING FROM,\nOUT OF OR IN CONNECTION WITH THE SOFTWARE OR THE USE OR OTHER DEALINGS IN\nTHE SOFTWARE.\n-->\n<style>\n/** Stylesheet for the puzzles */\n\n.sortable-code {\n    position: static;\n    padding-left: 0px;\n    margin-left: 2%;\n    float: left;\n    width: 94%;\n    max-width: 700px;\n}\n\n.sortable-code ul {\n    font-size: 90%;\n    font-family: monospace;\n    list-style: none;\n    background-color: #efefff;\n    padding-bottom: 10px;\n    padding-left: 0;\n    margin-left: 0;\n    border: 1px solid #efefff;;\n}\n.sortable-code ul:empty {\n  padding-bottom: 30px;\n}\n.sortable-code li, .sortable-code li:before, .sortable-code li:after {\n  box-sizing: content-box;\n}\nul.output {\n    background-color: #FFA;\n}\n.sortable-code li {\n    -moz-border-radius:10px;\n    -webkit-border-radius:10px;\n    border-radius: 10px;\n    background-color: white;\n    border:1px solid lightgray;\n    padding:4px;\n    padding-left:8px;\n    margin-top: 4px;\n    white-space: nowrap;\n    overflow: hidden;\n    cursor: move;\n}\n.sortable-code li:hover {\n    overflow: visible;\n}\n\n\n</style>\n\n<script>\nrequire(['jquery', 'jqueryui'], function($, jqui) { \n\n    function turn_on_touch() {\n        // Turn on touch functionality\n        // Detect touch support\n\n        /*!\n        * jQuery UI Touch Punch 0.2.3\n        *\n        * Copyright 2011\u20132014, Dave Furfero\n        * Dual licensed under the MIT or GPL Version 2 licenses.\n        *\n        * Depends:\n        *  jquery.ui.widget.js\n        *  jquery.ui.mouse.js\n        */\n        $.support.touch = 'ontouchend' in document;\n\n        // Ignore browsers without touch support\n        if (!$.support.touch) {\n            return;\n        }\n\n      var mouseProto = $.ui.mouse.prototype,\n          _mouseInit = mouseProto._mouseInit,\n          _mouseDestroy = mouseProto._mouseDestroy,\n          touchHandled;\n\n        /**\n        * Simulate a mouse event based on a corresponding touch event\n        * @param {Object} event A touch event\n        * @param {String} simulatedType The corresponding mouse event\n        */\n        function simulateMouseEvent (event, simulatedType) {\n\n        // Ignore multi-touch events\n        if (event.originalEvent.touches.length > 1) {\n          return;\n        }\n\n        event.preventDefault();\n\n        var touch = event.originalEvent.changedTouches[0],\n            simulatedEvent = document.createEvent('MouseEvents');\n\n        // Initialize the simulated mouse event using the touch event's coordinates\n        simulatedEvent.initMouseEvent(\n          simulatedType,    // type\n          true,             // bubbles                    \n          true,             // cancelable                 \n          window,           // view                       \n          1,                // detail                     \n          touch.screenX,    // screenX                    \n          touch.screenY,    // screenY                    \n          touch.clientX,    // clientX                    \n          touch.clientY,    // clientY                    \n          false,            // ctrlKey                    \n          false,            // altKey                     \n          false,            // shiftKey                   \n          false,            // metaKey                    \n          0,                // button                     \n          null              // relatedTarget              \n        );\n\n        // Dispatch the simulated event to the target element\n        event.target.dispatchEvent(simulatedEvent);\n        }\n\n        /**\n        * Handle the jQuery UI widget's touchstart events\n        * @param {Object} event The widget element's touchstart event\n        */\n        mouseProto._touchStart = function (event) {\n\n        var self = this;\n\n        // Ignore the event if another widget is already being handled\n        if (touchHandled || !self._mouseCapture(event.originalEvent.changedTouches[0])) {\n          return;\n        }\n\n        // Set the flag to prevent other widgets from inheriting the touch event\n        touchHandled = true;\n\n        // Track movement to determine if interaction was a click\n        self._touchMoved = false;\n\n        // Simulate the mouseover event\n        simulateMouseEvent(event, 'mouseover');\n\n        // Simulate the mousemove event\n        simulateMouseEvent(event, 'mousemove');\n\n        // Simulate the mousedown event\n        simulateMouseEvent(event, 'mousedown');\n        };\n\n        /**\n        * Handle the jQuery UI widget's touchmove events\n        * @param {Object} event The document's touchmove event\n        */\n        mouseProto._touchMove = function (event) {\n\n        // Ignore event if not handled\n        if (!touchHandled) {\n          return;\n        }\n\n        // Interaction was not a click\n        this._touchMoved = true;\n\n        // Simulate the mousemove event\n        simulateMouseEvent(event, 'mousemove');\n        };\n\n        /**\n        * Handle the jQuery UI widget's touchend events\n        * @param {Object} event The document's touchend event\n        */\n        mouseProto._touchEnd = function (event) {\n\n        // Ignore event if not handled\n        if (!touchHandled) {\n          return;\n        }\n\n        // Simulate the mouseup event\n        simulateMouseEvent(event, 'mouseup');\n\n        // Simulate the mouseout event\n        simulateMouseEvent(event, 'mouseout');\n\n        // If the touch interaction did not move, it should trigger a click\n        if (!this._touchMoved) {\n\n          // Simulate the click event\n          simulateMouseEvent(event, 'click');\n        }\n\n        // Unset the flag to allow other widgets to inherit the touch event\n        touchHandled = false;\n      };\n\n      /**\n       * A duck punch of the $.ui.mouse _mouseInit method to support touch events.\n       * This method extends the widget with bound touch event handlers that\n       * translate touch events to mouse events and pass them to the widget's\n       * original mouse event handling methods.\n       */\n      mouseProto._mouseInit = function () {\n        \n        var self = this;\n\n        // Delegate the touch handlers to the widget's element\n        self.element.bind({\n          touchstart: $.proxy(self, '_touchStart'),\n          touchmove: $.proxy(self, '_touchMove'),\n          touchend: $.proxy(self, '_touchEnd')\n        });\n\n        // Call the original $.ui.mouse init method\n        _mouseInit.call(self);\n      };\n\n      /**\n       * Remove the touch event handlers\n       */\n      mouseProto._mouseDestroy = function () {\n        \n        var self = this;\n\n        // Delegate the touch handlers to the widget's element\n        self.element.unbind({\n          touchstart: $.proxy(self, '_touchStart'),\n          touchmove: $.proxy(self, '_touchMove'),\n          touchend: $.proxy(self, '_touchEnd')\n        });\n\n        // Call the original $.ui.mouse destroy method\n        _mouseDestroy.call(self);\n      };\n    }; \n    \n    turn_on_touch(); // Run the jquery ui touch code.\n    \n    // =====================================================\n    //\n    // Now the code based on jsparsons (see licence at top)\n    // ----------------------------------------------------\n    // =====================================================\n\n    var MAX_INDENT = 7;\n    \n    // regexp used for trimming\n    var trimRegexp = /^\\s*(.*?)\\s*$/;\n    var translations = {\n        en: {\n            code_panel_label: 'Use drag and drop to reorder and indent the following code lines.',\n        }\n    };\n\n    // =================================================================\n    //\n    // ParsonsCodeLine object definition\n    // ---------------------------------\n    //\n    // Defines a line object skeleton with only code and indentation from\n    // a code string of the problem definition string (see parseCode)\n    // =================================================================\n    var ParsonsCodeLine = function(codestring, widget, allow_indent) {\n        this.widget = widget;\n        this.code = \"\";\n        this.indent = 0;\n        if (codestring) {\n            // Consecutive lines to be dragged as a single block of code have strings \"\\\\n\" to\n            // represent newlines => replace them with actual new line characters \"\\n\"\n            this.code = codestring.replace(/#distractor\\s*$/, \"\").replace(trimRegexp, \"$1\").replace(/\\\\n/g, \"\\n\");\n            if (allow_indent) {\n                this.indent = (codestring.length - codestring.replace(/^\\s+/, \"\").length) / 4;\n                this.indent = Math.min(this.indent, MAX_INDENT);\n            }\n            this.id = widget.id_prefix + ParsonsCodeLine.next_id++;\n        }\n    };\n\n    ParsonsCodeLine.next_id = 0; // Effectively a static variable of the 'class'.\n\n\n    // =================================================================\n    //\n    // ParsonsWidget object definition\n    // -------------------------------\n    // Creates a parsons widget.\n    // Init must be called after creating an object.\n    //\n    // ==================================================================\n    var ParsonsWidget = function(options) {\n        var defaults = {\n            'incorrectSound': false,\n            'x_indent': 35,\n            'feedback_cb': false,\n            'lang': 'en'\n        };\n\n        this.options = jQuery.extend({}, defaults, options);\n        this.id_prefix = options.sortableId + 'codeline';\n        if (translations.hasOwnProperty(this.options.lang)) {\n            this.translations = translations[this.options.lang];\n        } else {\n            this.translations = translations['en'];\n        }\n\n        // translate code_panel_label\n        if (!this.options.hasOwnProperty(\"code_panel_label\")) {\n            this.options.code_panel_label = this.translations.code_panel_label;\n        }\n    };\n\n    // Parse a problem definition given as a string and returns \n    // a list of ParsonCodeLines.\n    ParsonsWidget.prototype.parseCode = function(lines, allow_indent) {\n        var lineList = [],\n            lineObject,\n            that = this;\n        // Create line objects out of each codeline.\n        // Fields in line objects:\n        //   code: a string of the code, may include newline characters and \n        //     thus in fact represents a block of consecutive lines\n        //   indent: indentation level.\n        $.each(lines, function(index, item) {\n            lineObject = new ParsonsCodeLine(item, that, allow_indent);\n            // Initialize line object with code and indentation properties\n            if (lineObject.code.length > 0) {\n                // The line is non-empty, not just whitespace\n                lineList.push(lineObject);\n            }\n        });\n\n        return lineList;\n    };\n\n    ParsonsWidget.prototype.init = function(code_string) {\n        // Set up the widget, given the code to be displayed as a\n        // single string. \n        // Can't use jQuery to get element directly as it can't parse\n        // the element ID properly due to ___textareaId___ macro.\n        var lines = this.parseCode(code_string.split(\"\\n\"), true);\n        var that = this;\n        var sortableElement = document.getElementById(this.options.sortableId);\n        var sortable = $(sortableElement);\n        html = '<p>' + this.options.code_panel_label + '</p>' +\n            this.codeLinesToHTML(lines, this.options.sortableId);\n        \n        sortable.html(html);\n\n        this.manageDrags();\n    };\n\n\n    ParsonsWidget.prototype.calcIndent = function(leftDiff) {\n        // Return the indent based on the pixel difference given.\n        return Math.max(0, Math.round(leftDiff / this.options.x_indent));\n    };\n\n    ParsonsWidget.prototype.getCodeString = function(element_id) {\n        // Extract and return the code from the DOM as a single string.\n        var parent = document.getElementById(element_id);\n        var codeLines = parent.childNodes;\n        var code = \"\";\n        codeLines.forEach(item => {\n            var line = item.innerText.trim();\n            var indent = parseInt(item.style.marginLeft, 10);\n            indent = indent / this.options.x_indent;\n            var newLine = \" \".repeat(indent * 4) + line + '\\n';\n            code += newLine;\n        });\n        return code;\n    };\n\n    ParsonsWidget.prototype.getStudentCodeString = function() {\n        // Return the lines of code as a single string.\n        return this.getCodeString('ul-' + this.options.sortableId);\n    };\n\n    ParsonsWidget.prototype.displayError = function(message) {\n        if (this.options.incorrectSound && $.sound) {\n            $.sound.play(this.options.incorrectSound);\n        }\n        alert(message);\n    };\n\n    ParsonsWidget.prototype.setHTMLIndent = function(codelineID, indent) {\n        var codeLine = $(document.getElementById(codelineID));\n        codeLine.css(\"margin-left\", this.cssIndent(indent));\n    };\n\n    ParsonsWidget.prototype.cssIndent = function(indent) {\n        return this.options.x_indent * Math.min(MAX_INDENT, indent) + \"px\";\n    };\n\n    ParsonsWidget.prototype.codeLineToHTML = function(codeline) {\n        var style = codeline.indent ? ' style=\"margin-left:' + this.cssIndent(codeline.indent) + '\" ' : '';\n        return '<li id=\"' + codeline.id + '\" class=\"ace-highlight-code lang-py\"' + style + '>' + codeline.code + '<\\/li>';\n    };\n\n\n    ParsonsWidget.prototype.codeLinesToHTML = function(codeLines, destinationID) {\n        // Construct and return a single <ul> element with the HTML\n        // for the given list of lines as embedded LI elements. The UL\n        // element's ID is destinationID.\n        var lineHTML = [];\n        codeLines.forEach(line => {\n            lineHTML.push(this.codeLineToHTML(line));\n        });\n        return '<ul id=\"ul-' + destinationID + '\">' + lineHTML.join('') + '</ul>';\n    };\n\n    ParsonsWidget.prototype.manageDrags = function() {\n        // Set up the code pane as a sortable UI elements with drag\n        // management for indentation. An annoying issue here is that\n        // the ___textareaId___ macro produces IDs that jQuery cannot\n        // handle, so we have to use getElementById everywhere.\n        var that = this;\n        var sortableElement = document.getElementById('ul-' + this.options.sortableId);\n        var sortable = $(sortableElement).sortable({\n            start: function(event, ui) {\n             },\n            stop: function(event, ui) {\n                if ($(event.target)[0] != ui.item.parent()[0]) {\n                    return;\n                }\n                // The helper isn't where it seems to be because of the margin!\n                var currentMargin = parseInt(ui.item.css('margin-left'));\n                var realPosition = currentMargin + ui.position.left - ui.item.parent().position().left;\n                var indent = that.calcIndent(realPosition);\n                that.setHTMLIndent(ui.item[0].id, indent);\n            },\n            grid: [this.options.x_indent, 1]\n        });\n        sortable.addClass(\"output\");\n    };\n    \n   window.ParsonsWidget = ParsonsWidget;\n   \n    function displayErrors(fb) {\n        if(fb.errors.length > 0) {\n            alert(fb.errors[0]);\n        }\n    } \n\n    $(document).ready(function() {\n        var sortableId = 'sortable-___textareaId___';\n        var crTextareaId = 'CR-parsons-code-___textareaId___';\n        var parson = new ParsonsWidget({\n            'sortableId': sortableId\n        });\n\n        var codeTextArea = document.getElementById(crTextareaId);\n\n        var code = codeTextArea.value; // Extract the code as a single multiline string.\n        parson.init(code);\n        \n        function sync_to_coderunner(mutations) {\n            codeTextArea.value = parson.getStudentCodeString();\n        };\n        dropArea=document.getElementById('ul-' + sortableId);\n        let observer=new MutationObserver(sync_to_coderunner);\n        let observerOptions = {\n            childList: true,\n            attributes: true,\n            characterData: false,\n            subtree: false,\n            attributeFilter: ['style'],\n            attributeOldValue: false,\n            characterDataOldValue: false\n        };\n        observer.observe(dropArea, observerOptions);\n        if (window.ace && window.applyAceHighlighting) {\n            window.applyAceHighlighting(window.ace, dropArea);\n        }\n    });\n});\n</script>\n\n<html>\n    <head>\n        <title>Simple js-parsons example</title>\n    </head>\n    <body>\n        <div id=\"sortable-___textareaId___\" class=\"sortable-code\">\n        </div>\n        <div style=\"clear:both;\"></div>\n        <div id=\"CRactual\"  style=\"display: none;\">\n        <h2>Code Here:</h2>\n<textarea id=\"CR-parsons-code-___textareaId___\" name='CR-parsons-code' class=\"coderunner-ui-element\" rows=\"10\" cols=\"40\">\n{{ jumbled_code }}\n</textarea>\n        </div>\n    </body>\n</html>\n\n{% endmacro %}\n"}
]]></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin>html</uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
    </testcases>
  </question>

</quiz>